# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
```

## #01 - 

How many columns does this file contain?

```python
hhcaps.shape
```

Output:
```
(12000, 39)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
answer = 0
## Solution goes here
my_df = hhcaps.columns
answer = my_df.shape[0]
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

```python

```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

## Solution goes here
answer = len(hhcaps['State'].unique())
```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

```python

```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
answer = ''
## Solution goes here
hhcaps_state = hhcaps.groupby(['State'])
answer = hhcaps_state['State'].count().idxmax()
```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

```python

```

```python

```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
answer = ''
## Solution goes here
answer = hhcaps_state['Star Rating for health team communicated well with them'].mean().idxmax()
```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

```python

```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
answer = 0.0

## Solution goes here
answer = hhcaps[hhcaps['Type of Ownership'] == 'Hospital Based Program']['Star Rating for how patients rated overall care from agency'].mean()
```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
assert(round(answer,5) == round(3.6925207756232687,5))

```

```python

```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

<!-- grade_id: q06-solution, points: 0, grade: false, solution: true -->
```python
answer = ''
for column in hhcaps.columns:
    if len(hhcaps.groupby([column]).count()) == 1:
        answer = column
answer
```

Output:
```
'Offers Nursing Care Services'
```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())

```

```python

```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

#everything = hhcaps[(hhcaps['Offers Nursing Care Services']== True)&(hhcaps['Offers Physical Therapy Services']== True)&(hhcaps['Offers Occupational Therapy Services']== True)&(hhcaps['Offers Speech Pathology Services']== True)&(hhcaps['Offers Medical Social Services']== True)&(hhcaps['Offers Home Health Aide Services']== True)]
## Solution goes here
offers = []
for colum in hhcaps.columns:
    if 'Offers' in colum:
        offers.append(colum)
h_offer = hhcaps
for offer in offers:
    h_offer = h_offer[h_offer[offer]== True]
answer = len(h_offer)

```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)

```

```python

```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


<!-- grade_id: q08-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

## Solution goes here
hhcaps_pn = hhcaps[hhcaps['Number of completed Surveys'] != 'Not Available']
answer = hhcaps_pn['Number of completed Surveys'].astype(float).median()
answer
```

Output:
```
84.0
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer == 84)

```

```python

```

```python

```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

<!-- grade_id: q09-solution, points: 0, grade: false, solution: true -->
```python
answer = 0

## Solution goes here
hhcaps_sl = hhcaps[hhcaps['City'].str.strip() == 'SAINT LOUIS']
answer = len(hhcaps_sl['Provider Name'].unique())
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)

```

```python

```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python
answer = 0
field = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'

## Solution goes here
hh = hhcaps[hhcaps[field] != 'Not Available']
hh[field] = hh[field].astype(float)
answer= hh.groupby(['State']).mean()[field].idxmax()
```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/ipykernel_launcher.py:6: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: http://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```

```python

```